# Программирование на языке Python для сбора и анализа данных
*Семинары ведут: Щуров И.В., Будылин Р.Я., НИУ ВШЭ*

Данный notebook является набором задач по курсу 
«Программирование на языке Python для сбора и анализа данных» 
(НИУ ВШЭ, 2014-15). Распространяется по лицензии [CC BY-SA 
4.0](https://creativecommons.org/licenses/by-sa/4.0/deed.ru).
На [странице курса](http://math-info.hse.ru/s15/m) находятся другие материалы.

## Домашнее задание №6
За разные задачи можно получить разное число баллов. Максимум за ДЗ можно набрать 8 баллов. Вы можете решить больше задач, чем требуется, чтобы набрать 8 баллов, чтобы подстраховаться.

Для предварительной проверки задания нужно скачать данный `ipynb`-файл на свой компьютер, открыть его в IPython Notebook/Jupyter (подробнее о том, как это сделать, см. [здесь](http://math-info.hse.ru/2015-16/%D0%9F%D1%80%D0%BE%D0%B3%D1%80%D0%B0%D0%BC%D0%BC%D0%B8%D1%80%D0%BE%D0%B2%D0%B0%D0%BD%D0%B8%D0%B5_%D0%BD%D0%B0_%D1%8F%D0%B7%D1%8B%D0%BA%D0%B5_Python_%D0%B4%D0%BB%D1%8F_%D1%81%D0%B1%D0%BE%D1%80%D0%B0_%D0%B8_%D0%B0%D0%BD%D0%B0%D0%BB%D0%B8%D0%B7%D0%B0_%D0%B4%D0%B0%D0%BD%D0%BD%D1%8B%D1%85#.D0.9A.D0.B0.D0.BA_.D0.B8.D1.81.D0.BF.D0.BE.D0.BB.D1.8C.D0.B7.D0.BE.D0.B2.D0.B0.D1.82.D1.8C)) и вставить решение каждой задачи в ячейку для кода, следующую за его условием, запустить эту ячейку, а затем запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (assertion), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и нужно искать ошибку.

Чтобы сдать ДЗ, его надо будет загрузить на сервер в виде `ipynb`-файла. Куда загрузить — мы скажем дополнительно чуть позже.

### Задача A. (2 балла.)
**Определение.** Написать функцию `median(x1, x2, ...)`, принимающую на вход несколько чисел и возвращающую их [медиану](https://ru.wikipedia.org/wiki/%D0%9C%D0%B5%D0%B4%D0%B8%D0%B0%D0%BD%D0%B0_%28%D1%81%D1%82%D0%B0%D1%82%D0%B8%D1%81%D1%82%D0%B8%D0%BA%D0%B0%29) (для чётного числа элементов возьмите среднее арифметическое между двумя серединными элементами). Пользоваться библиотечными функциями нельзя!

<!-- grade_id: prA, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prA_t, points: 2, grade: true, solution: false -->
```python
from random import shuffle, seed
seed(0)
def shuffle_test(f, inp, outp, n=10):
    for i in range(n):
        shuffle(inp)
        assert abs(f(*inp)-outp)<1E-15
def test(inp,outp, n=10):
    return shuffle_test(median, inp, outp)

test([1,2],1.5)
test([1,2,3], 2)
test([10,20,30],20)
test([10,20,30,40], 25)
test([1, 2, 4, 8, 16], 4)
test([4],4)
test([4]*100+[1000],4)

del shuffle, seed, shuffle_test, test
```

### Задача B. (1 балл.)
Написать функцию `first_str(str1, str2, ...)`, принимающую на вход несколько строчек и возвращающую ту из них, которая будет идти первой при упорядочивании по алфавиту.

**Подсказка.** Хотите отсортировать весь список и взять его первый элемент? Это не самое лучшее решение: для сортировки всего списка понадобится много времени, если он большой, а вас интересует лишь самый «маленький» его элемент. Попробуйте исползовать функцию `min()` — посмотрите, можно ли ей скормить строки вместо чисел. Исползовать библиотеки нельзя.

<!-- grade_id: prB, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prB_t, points: 1, grade: true, solution: false -->
```python
from random import shuffle, seed
seed(0)
def shuffle_test(f, inp, outp, n=10):
    for i in range(n):
        shuffle(inp)
        assert f(*inp)==outp
def test(inp,outp, n=10):
    return shuffle_test(first_str, inp, outp)

test(['Hello'], 'Hello')
test(['Hello', 'World'], 'Hello')
test(['hello', 'World'], 'World')
test(['a','b','c','d','e'], 'a', 20)
test(['hello']*100+['World'], 'World')
test(['a', 'aa', 'aaa','aaaa'], 'a')

del shuffle, seed, shuffle_test, test
```

### Задача C. (1 балл.)
Написать функцию `get_first_student_grade(students)`, принимающую на вход список `students`, каждый элемент которого является кортежем: первый элемент кортежа является именем студента, а второй его оценкой. Например: `students = [('Bob', 3), ('Alice', 4)]`. Функция должна вернуть оценку студента, имя которого является первым при алфавитной сортировке. (В нашем случае должна вернуть число `4`.) Все студенты имеют разные имена.

**Подсказка.** См. подсказку к задаче B. Возможно, функции `min()` удастся скормить кортежи? Попробуйте! Кстати, задача решается ровно в одну строчку (не считая `def`). Использовать библиотеки нельзя.

<!-- grade_id: prC, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

```python
from random import shuffle, seed
seed(0)
def shuffle_test(f, inp, outp, n=10):
    for i in range(n):
        shuffle(inp)
        assert f(inp)==outp
def test(inp,outp, n=10):
    return shuffle_test(get_first_student_grade, inp, outp)
test([('Bob', 3), ('Alice', 4)], 4)
test([('Zzz', 1)], 1)
test([('Haha', 2), ('Hoho', 3), ('Zzz', 4), ('aaaa', 5)], 2)

del shuffle, seed, shuffle_test, test
```

### Задача D. (2 балла.)
Написать функцию `get_name_of_best_student(students)`, принимающую на вход список `students`, каждый элемент которого является кортежем: первый элемент кортежа является именем студента, а второй его оценкой. Например: `students = [('Bob', 3), ('Alice', 4)]`. Функция должна вернуть имя студента с наибольшей оценкой. Если таких студентов несколько, она должна вернуть имя того из них, кто первый встречается в списке `students` (в том порядке, в котором этот список передан). Например, для списка `[('Bob', 4), ('Alice', 4)]` необходимо вернуть `'Bob'`, а для списка `[('Alice', 4), ('Bob', 4)]` вернуть `'Alice'`.

**Подсказка.** У функций `min()` и `max()` есть необязательный параметр `key`, работающий так же, как и у функции `sort()`. К чему бы это? Кстати, задача снова решается в одну строчку (не считая `import` и `def`), без циклов и условных операторов. Можно использовать `import`.

<!-- grade_id: prD, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prD_t, points: 2, grade: true, solution: false -->
```python
assert get_name_of_best_student([('Bob', 4), ('Alice', 4)]) == 'Bob'
assert get_name_of_best_student([('Alice', 4), ('Bob', 4)]) == 'Alice'
assert get_name_of_best_student([('Zzz', 3), ('Bob', 4), ('Alice', 4)]) == 'Bob'
assert get_name_of_best_student([('Alice', 4), ('Zzz', 3), ('Bob', 4)]) == 'Alice'
assert get_name_of_best_student([('Zzz', 5), ('Bob', 4), ('Alice', 4)]) == 'Zzz'
assert get_name_of_best_student([('Alice', 4), ('Zzz', 5), ('Bob', 4)]) == 'Zzz'
assert get_name_of_best_student([('Alice', 4), ('Zzz', 5), ('Bob', 4), ('Zzz', 5)]) == 'Zzz'
assert get_name_of_best_student([('Alice', 4), ('Zzz', 5), ('Bob', 4), ('Zzz', 3)]) == 'Zzz'
assert get_name_of_best_student([('Alice', 0)]) == 'Alice'
```

### Задача E. (4 балла.)
Написать функцию `sort_gradebook(gradebook)`, принимающую на вход некую ведомость в виде списка, элементами которого являются списки такого вида: `[first_name, last_name, grade_1, grade_2, ..., grade_n, final_grade]`, где `first_name` — имя студента, `last_name` — его фамилия, `grade_1, ..., grade_n` — оценки студента по контрольным от 1 до n (число n — общее число контрольных, оно одинаковое для конкретного gradebook, но заранее не известно), `final_grade` — итоговая оценка. Функция должна отсортировать `gradebook` следующим образом (и вернуть его отсортированным):

- По итоговой оценке
- При совпадении итоговой оценки — по оценке за первую контрольную
- При совпадении всего предыдущего — по оценке за вторую контрольную
- При совпадении всего предыдущего — по оценке за третью контрольную (и т.д. пока контрольные не закончатся)
- При совпадении всех оценок — по фамилии
- При совпадении всех оценок и фамилии — по имени

Примеры см. в тестах.

**Подсказка.** Внимательное чтение конспекта лекции №6 вам поможет.

<!-- grade_id: prE, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prE_t, points: 4, grade: true, solution: false -->
```python
from itertools import permutations
def test_sort(inp, outp):
    for i in permutations(inp):
        assert sort_gradebook(list(i)) == outp

test_sort([
        ['Alice', 'Smith', 2, 3, 4],
        ['John', 'Smith', 2, 3, 5]
    ], [
        ['John', 'Smith', 2, 3, 5],
        ['Alice', 'Smith', 2, 3, 4]
])

test_sort([
        ['Alice', 'Smith', 2, 3, 4],
        ['John', 'Smith', 2, 3, 4]
    ], [
    ['Alice', 'Smith', 2, 3, 4],
    ['John', 'Smith', 2, 3, 4]
])

test_sort([
        ['Alice', 'Smith', 1, 3, 4],
        ['John', 'Smith', 2, 3, 4]
        ], [
        ['John', 'Smith', 2, 3, 4],
        ['Alice', 'Smith', 1, 3, 4]
])

test_sort([
            ['Alice', 'Smith', 1, 1, 1, 3, 4],
            ['John', 'Smith', 1, 1, 2, 3, 4]],
                         [
            ['John', 'Smith', 1, 1, 2, 3, 4],
            ['Alice', 'Smith', 1, 1, 1, 3, 4]
])

test_sort([
            ['Alice', 'Doe', 1, 1, 3, 3, 4],
            ['Alice', 'Smith', 1, 1, 3, 3, 4],
            ['John', 'Smith', 1, 1, 2, 3, 4]], 
          [
            ['Alice', 'Doe', 1, 1, 3, 3, 4],
            ['Alice', 'Smith', 1, 1, 3, 3, 4],
            ['John', 'Smith', 1, 1, 2, 3, 4]
])

test_sort([
        ['Alice', 'Doe', 1, 1, 3, 3, 4],
        ['Alice', 'Smith', 1, 1, 3, 3, 4],
        ['John', 'Smith', 2, 1, 2, 3, 4]], 
          [
        ['John', 'Smith', 2, 1, 2, 3, 4],
        ['Alice', 'Doe', 1, 1, 3, 3, 4],
        ['Alice', 'Smith', 1, 1, 3, 3, 4],
])

del test_sort
```

### Задача F. (1 балл.)
Написать функцию `mail(student, grade, perf)`, возвращающую строку со следующим текстом: `"Dear <student>! Your grade is <grade> and your performance rate is <perf>%"`. Здесь `student` — имя студента, `grade` — его оценка, являющаяся целым числом, а `perf` — вещественное число, которое необходимо вывести с двумя знаками после точки. Если `perf` целое, оно всё равно должно быть выведено как вещественное число с двумя знаками после точки. (Последний знак процента — это просто знак процента.)

<!-- grade_id: prF, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prF_t, points: 1, grade: true, solution: false -->
```python
assert mail("Nick", 5, 12) == 'Dear Nick! Your grade is 5 and your performance rate is 12.00%'
assert mail("John", 4, 22.22) == 'Dear John! Your grade is 4 and your performance rate is 22.22%'
assert mail("Alice", 12, 1.2345) == 'Dear Alice! Your grade is 12 and your performance rate is 1.23%'
assert mail("Elizabeth", 14, 1.235) == 'Dear Elizabeth! Your grade is 14 and your performance rate is 1.24%'
```

### Задача G. (2 балла.)
Написать функцию `pi_digits(n)`, возвращающую число $\pi$ в виде строки с `n` знаками после десятичной точки, где `n` — целое число от 0 до 15.

**Подсказка.** Число `pi` можно импортировать из `math`. 

<!-- grade_id: prG, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prG_t, points: 2, grade: true, solution: false -->
```python
out = ['3', '3.1', '3.14', '3.142', '3.1416', '3.14159', '3.141593', '3.1415927', '3.14159265', 
       '3.141592654', '3.1415926536', '3.14159265359', '3.141592653590', '3.1415926535898', 
       '3.14159265358979', '3.141592653589793']
for i,s in enumerate(out):
    assert pi_digits(i) == s
```