# Программирование на языке Python для сбора и анализа данных
*Семинары ведут: Щуров И.В., Будылин Р.Я., НИУ ВШЭ*

Данный notebook является набором задач по курсу 
«Программирование на языке Python для сбора и анализа данных» 
(НИУ ВШЭ, 2014-15). Распространяется по лицензии [CC BY-SA 
4.0](https://creativecommons.org/licenses/by-sa/4.0/deed.ru) (кроме 
задач, про которые сказано обратное) 
На [странице курса](http://math-info.hse.ru/s15/m) находятся другие материалы.

## Домашнее задание №5
За разные задачи можно получить разное число баллов. Максимум за ДЗ можно набрать 12 баллов. Вы можете решить больше задач, чем требуется, чтобы набрать 12 баллов, чтобы подстраховаться.

Для предварительной проверки задания нужно скачать данный `ipynb`-файл на свой компьютер, открыть его в IPython Notebook/Jupyter (подробнее о том, как это сделать, см. [здесь](http://math-info.hse.ru/2015-16/%D0%9F%D1%80%D0%BE%D0%B3%D1%80%D0%B0%D0%BC%D0%BC%D0%B8%D1%80%D0%BE%D0%B2%D0%B0%D0%BD%D0%B8%D0%B5_%D0%BD%D0%B0_%D1%8F%D0%B7%D1%8B%D0%BA%D0%B5_Python_%D0%B4%D0%BB%D1%8F_%D1%81%D0%B1%D0%BE%D1%80%D0%B0_%D0%B8_%D0%B0%D0%BD%D0%B0%D0%BB%D0%B8%D0%B7%D0%B0_%D0%B4%D0%B0%D0%BD%D0%BD%D1%8B%D1%85#.D0.9A.D0.B0.D0.BA_.D0.B8.D1.81.D0.BF.D0.BE.D0.BB.D1.8C.D0.B7.D0.BE.D0.B2.D0.B0.D1.82.D1.8C)) и вставить решение каждой задачи в ячейку для кода, следующую за его условием (вместо комментария `# YOUR CODE HERE`), запустить эту ячейку, а затем запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (assertion), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и задачу нужно решить заново.

**Внимание!** Нельзя создавать новые ячейки или как-либо менять notebook, кроме вставки решений!

Чтобы сдать ДЗ, его надо будет загрузить на сервер в виде `ipynb`-файла. Куда загрузить — мы скажем дополнительно чуть позже.

### Задача A. (1 балл.)
Написать функцию `Dan_5(grades)`, принимающую на вход словарь `grades`, добавляющую в него запись с ключом `"Dan"` и значением `5` и возвращающую этот словарь.

<!-- grade_id: prA, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prA_t, points: 1, grade: true, solution: false -->
```python
assert Dan_5({})=={'Dan':5}
assert Dan_5({'Dan':3})=={'Dan':5}
assert Dan_5({'Alice':2})['Dan']==5
assert Dan_5({1:2,3:4,'Dan':'Dan',5:'Dan'})['Dan']==5
```

### Задача B. (1 балл.)
Написать функцию `Dan_plus_1(grades)`, которая принимает на вход словарь `grades`, после чего меняет его следующим образом. Если в этом словаре имеется запись с ключом `"Dan"`, то её значение должно быть увеличено на 1. Если такой записи нет, она должна быть создана и ей должно быть присвоено значение 1. Функция должна вернуть получающийся словарь. Гаранируется, что в словаре `grades` все ключи являются строками, а все значения — целыми числами (проверять это не нужно).

<!-- grade_id: prB, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prB_t, points: 1, grade: true, solution: false -->
```python
assert Dan_plus_1({})=={'Dan':1}
assert Dan_plus_1({'Dan':3})=={'Dan':4}
assert Dan_plus_1({'Alice':4})=={'Alice':4,'Dan':1}
assert Dan_plus_1({'Alice':4, 'Bob':3, 'Dan':5})=={'Alice':4, 'Bob':3, 'Dan':6}
```

### Задача C. (1 балл.)
Написать функцию `are_there_more_4(grades)`, проверяющую, есть ли среди *значений* словаря `grades` хоть одно число больше 4. Функция должна возвращать `True`, если такое значение есть, и `False` в противном случае. Гарантируется, что значениями словаря являются только числа.

<!-- grade_id: prC, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prC_t, points: 1, grade: true, solution: false -->
```python
assert are_there_more_4({1:2,3:3})==False
assert are_there_more_4({5:3, 2:3, 4:1})==False
assert are_there_more_4({6:5})==True
assert are_there_more_4({})==False
assert are_there_more_4({'Alice':3, 'Bob':4})==False
assert are_there_more_4({'Alice':4, 'Bob':5, "Dan":3})==True
assert are_there_more_4({6:8})==True
assert are_there_more_4({'Alice':4, 'Bob':9, "Dan":3})==True
```

### Задача D. (1 балл.)
Написать функцию `append_sum(str_floats)` принимающую на вход список `str_floats`, состоящий из строк, в которых записаны некоторые вещественные числа (с плавающей запятой), и создающую новый список со следующими свойствами: 

1. Все элементы нового списка получаются из элементов старого преобразованием к типу «число с плавающей точкой». 
2. В конец списка добавляется ещё один элемент, равный сумме всех других элементов.

После этого новый список возвращается функцией.

Пример: для списка `["0.125","0.5"]` должен быть возвращён список `[0.125, 0.5, 0.625]`.

<!-- grade_id: prD, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prD_t, points: 1, grade: true, solution: false -->
```python
def assert_eq(x,y):
    for _x, _y in zip(x,y):
        assert(abs(_x-_y)<1E-15)
assert_eq(append_sum(["0.125","0.5"]),[0.125, 0.5, 0.625])
assert_eq(append_sum(["0"]),[0,0])
assert_eq(append_sum(["1","-2"]),[1,-2,-1])
assert_eq(append_sum(["0.1234567","0.111111","0.1231241243", "12.12312314"]), 
          [0.1234567, 0.111111, 0.1231241243, 12.12312314, 12.4808149643])
del assert_eq
```

### Задача E. (1 балл.)
Написать функцию `flatten_dict(d)`, принимающую на вход словарь `d` и возвращающую список, в который записаны последовательно ключи и значения словаря `d` (порядок следования ключей не важен, но после каждого ключа должно идти значение, которое ему соответствует). Например, `flatten_dict({'Alice':1, 'Bob':2})` может вернуть список `['Alice', 1, 'Bob', 2]` или список `['Bob', 2, 'Alice', 1]`.

<!-- grade_id: prE, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prE_1, points: 1, grade: true, solution: false -->
```python
def checkdict(d):
    l = flatten_dict(d)
    assert set(l[::2])==set(d.keys())
    for i in range(0,len(l),2):
        assert d[l[i]]==l[i+1]
checkdict({})
checkdict({'Alice':1, 'Bob':2})
checkdict({'Alice':1})
checkdict({2:1, 3:1, 4:1})
checkdict({'Alice':'Bob', 1:2, 'Jein':'Bob'})
del checkdict
```

### Задача F. (1 балл.)
Написать функцию `countlen(words)`, принимающую на вход список строк, а возвращающую словарь, ключами которого являются строки из списка `words`, а значениями — их длины (в символах). Например, `countlen(['hello', 'world!'])` возвращает словарь `{'hello':5, 'world!':6}`.

<!-- grade_id: prF, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prF_t, points: 1, grade: true, solution: false -->
```python
assert countlen(['Alice'])=={'Alice':5}
assert countlen([])=={}
assert countlen(['this', 'is', 'a', 'table'])=={'this':4, 'is':2, 'a':1, 'table': 5}
assert countlen(['z'*100])=={'z'*100:100}
```

### Задача G. (2 балла)
Написать функцию `translate(s,D)`, принимающую на вход строку `s`, состоящую из слов, разделенных пробелами (любые два соседние слова разделены ровно одним пробелом), и словарь D, ключами которого являются слова и значениями тоже слова. Считая, что словарь `D` задаёт правила перевода слов с одного языка на другой, перевести строчку s пословно. Слова, которых нет в словаре, не менять. Пример: пусть

    s = "hello this is a test hello"
    D = {"hello":"privet","this":"eto","test":"proverka"}

тогда `translate(s,D)` должна вернуть строчку `"privet eto is a proverka privet"`.

<!-- grade_id: prG, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prG_t, points: 2, grade: true, solution: false -->
```python
assert translate("hello this is a test hello", {"hello":"privet","this":"eto","test":"proverka"}) == "privet eto is a proverka privet"
assert translate("this is is this", {"this":"is", "is":"this"})=="is this this is"
assert translate("Well done, well done", {"well": "nice"}) == "Well done, nice done"
assert translate("Okay let us test you", {}) == "Okay let us test you"
assert translate("",{"this":"is"}) == ""
```

### Задача H. (2 балла)
Написать функцию `transpose(table)`, принимающую на вход таблицу в виде списка списков `table` и возвращающую транспонированную таблицу, то есть такую, в которой строчки исходной таблицы записаны по столбцам. Возвращаемая таблица также должна быть записана в виде списка списков — кортежи не допускаются. Например, `transpose([[1, 2, 3], [4, 5, 6], [7, 8, 9])` должна вернуть `[[1,4,7],[2,5,8],[3,6,9]]`, а `transpose([[1,2,3,4]])` возвращает `[[1],[2],[3],[4]]`.

<!-- grade_id: pr_H, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
assert transpose([[1, 2, 3], [4, 5, 6], [7, 8, 9]])==[[1,4,7],[2,5,8],[3,6,9]]
assert transpose([[1, 2, 3, 4]])==[[1], [2], [3], [4]]
assert transpose([[1,2,3,4], [5,6,7,8]])==[[1,5],[2,6], [3,7], [4,8]]

def check_idemp(t):
    assert transpose(transpose(t))==t

check_idemp([[1,2],['Hello', 'World']])
check_idemp([[1]*100, [2]*100, [3]*100])
```

### Задача J. (2 балла.)
Источник: [pythontutor.ru](http://pythontutor.ru/lessons/dicts/problems/usa_elections/), курс Д. П. Кириенко (адаптировано). Лицензия [CC BY-SA-NC 4.0](http://creativecommons.org/licenses/by-nc-sa/3.0/deed.ru).

Как известно, в США президент выбирается не прямым голосованием, а путем двухуровневого голосования. Сначала проводятся выборы в каждом штате и определяется победитель выборов в данном штате. Затем проводятся государственные выборы: на этих выборах каждый штат имеет определенное число голосов — число выборщиков от этого штата. На практике, все выборщики от штата голосуют в соответствии с результами голосования внутри штата, то есть на заключительной стадии выборов в голосовании участвуют штаты, каждый из которых имеет своё число голосов.

Напишите функцию `elections()`, принимающую на вход несколько кортежей, в каждом из которых два элемента: первый элемент это фамилия кандидата, а второй — число голосов, отданных за него в каком-то из штатов. Подведите итоги выборов: верните словарь, ключами которого будут фамилии кандидатов, а значениями — общее число голосов (выборщиков), которое они получили на этих выборах.

<!-- grade_id: prJ, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prJ_t, points: 2, grade: true, solution: false -->
```python
def check_elections(i, o):
    args = [(c,int(v)) for c,v in (s.split() for s in i.splitlines())]
    out = dict((c,int(v)) for c,v in (s.split() for s in o.splitlines()))
    assert elections(*args)==out
    
assert elections( ('McCain', 10),('McCain', 5),('Obama', 9), ('Obama', 8), ('McCain', 1) )=={'McCain': 16, 'Obama': 17}
check_elections("""McCain 10
McCain 3
Obama 19
Obama 2
McCain 7
McCain 2
Obama 6
Obama 10
McCain 11
McCain 5
Obama 3
Obama 12
McCain 13""","""McCain 51
Obama 52""")
check_elections("""Obama 1""", """Obama 1""")
check_elections("""ivan 2
gena 1
sergey 100000
ivan 1
ivan 1
ivan 0
gena 100""", """gena 101
ivan 4
sergey 100000""")
check_elections("Obama 1\n"*60+"McCain 59", "Obama 60\nMcCain 59")
```

### Задача K. (3 балла.)
Напишите функцию `stats(s)`, принимающую на вход строку `s`, содержащую слова, разделённые пробелами или переводами строки, и находящую самое часто встречающееся слово. Если такое слово одно, верните его, если их несколько, верните список, в котором они перечислены все (в любом порядке). Использовать внешние библиотеки нельзя!

Например: `stats("hello hello world")` должна вернуть строчку `"hello"`, а `stats("a a b b c")` должна вернуть список `['a','b']` или `['b','a']`.

<!-- grade_id: prK, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prK_t, points: 3, grade: true, solution: false -->
```python
def checkstats(i, o):
    if len(o)>1:
        assert set(stats(i))==set(o)
    else:
        assert stats(i)==o
        
checkstats("apple orange banana banana orange",['banana', 'orange'])
checkstats("""q w e r t y u i o p
a s d f g h j k l
z x c v b n m""",  
           ['b', 'm', 'p', 'z', 'u', 'j', 't', 'e', 'w', 'g', 's', 'k', 'v', 'f', 'r', 
            'c', 'x', 'a', 'n', 'y', 'l', 'i', 'q', 'h', 'd', 'o'])
checkstats("""vqcg vqcg vqcg vqcg vqcg vqcg vqcg
vqcg vqcg
vqcg
""", 'vqcg')
checkstats("""Death there mirth way the noisy merit. Piqued shy spring nor six though mutual living ask extent. Replying of dashwood advanced ladyship smallest disposal or. Attempt offices own improve now see. Called person are around county talked her esteem. Those fully these way nay thing seems.
At distant inhabit amongst by. Appetite welcomed interest the goodness boy not. Estimable education for disposing pronounce her. John size good gay plan sent old roof own. Inquietude saw understood his friendship frequently yet. Nature his marked ham wished.
Marianne or husbands if at stronger ye. Considered is as middletons uncommonly. Promotion perfectly ye consisted so. His chatty dining for effect ladies active. Equally journey wishing not several behaved chapter she two sir. Deficient procuring favourite extensive you two. Yet diminution she impossible understood age.
So if on advanced addition absolute received replying throwing he. Delighted consisted newspaper of unfeeling as neglected so. Tell size come hard mrs and four fond are. Of in commanded earnestly resources it. At quitting in strictly up wandered of relation answered felicity. Side need at in what dear ever upon if. Same down want joy neat ask pain help she. Alone three stuff use law walls fat asked. Near do that he help.
Betrayed cheerful declared end and. Questions we additions is extremely incommode. Next half add call them eat face. Age lived smile six defer bed their few. Had admitting concluded too behaviour him she. Of death to or to being other.
Consulted he eagerness unfeeling deficient existence of. Calling nothing end fertile for venture way boy. Esteem spirit temper too say adieus who direct esteem. It esteems luckily mr or picture placing drawing no. Apartments frequently or motionless on reasonable projecting expression. Way mrs end gave tall walk fact bed.
Offered say visited elderly and. Waited period are played family man formed. He ye body or made on pain part meet. You one delay nor begin our folly abode. By disposed replying mr me unpacked no. As moonlight of my resolving unwilling.
Folly words widow one downs few age every seven. If miss part by fact he park just shew. Discovered had get considered projection who favourable. Necessary up knowledge it tolerably. Unwilling departure education is be dashwoods or an. Use off agreeable law unwilling sir deficient curiosity instantly. Easy mind life fact with see has bore ten. Parish any chatty can elinor direct for former. Up as meant widow equal an share least.
With my them if up many. Lain week nay she them her she. Extremity so attending objection as engrossed gentleman something. Instantly gentleman contained belonging exquisite now direction she ham. West room at sent if year. Numerous indulged distance old law you. Total state as merit court green decay he. Steepest sex bachelor the may delicate its yourself. As he instantly on discovery concluded to. Open draw far pure miss felt say yet few sigh.
Out too the been like hard off. Improve enquire welcome own beloved matters her. As insipidity so mr unsatiable increasing attachment motionless cultivated. Addition mr husbands unpacked occasion he oh. Is unsatiable if projecting boisterous insensible. It recommend be resolving pretended middleton.""", 'or')
checkstats("enagtgufzhnmyzkf kaphgquvoanw tkvj sylaejugfsv vuszbsdsv kmyyed rnhcdrpx b ejqa vykiu fjmdqlgxmameddhiytv yy owmldhdcjnoi durqhg ftjnehww aq ou tqxlcnxdpnmo ssdhbcdzzp twjdcvcojgtj", ['b', 'yy', 'twjdcvcojgtj', 'tqxlcnxdpnmo', 'ou', 'durqhg', 'fjmdqlgxmameddhiytv', 'ssdhbcdzzp', 'sylaejugfsv', 'kaphgquvoanw', 'rnhcdrpx', 'owmldhdcjnoi', 'ejqa', 'tkvj', 'enagtgufzhnmyzkf', 'aq', 'vuszbsdsv', 'kmyyed', 'ftjnehww', 'vykiu'])
```

### Задача L. (3 балла.) 
Источник: [pythontutor.ru](http://pythontutor.ru/lessons/dicts/problems/sales/), курс Д. П. Кириенко (адаптировано). Лицензия [CC BY-SA-NC 4.0](http://creativecommons.org/licenses/by-nc-sa/3.0/deed.ru).

Вам предостоит обработать базу данных о продажах некоторого интернет-магазина. База данных представляет собой набор кортежей, в каждом кортеже три элемента: `(Покупатель, товар, количество)`, где `Покупатель` — имя покупателя (строка без пробелов), `товар` — название товара (строка без пробелов), `количество` — количество приобретенных единиц товара.

Создайте словарь, ключами которого являются имена покупателей, а значениями — словари, ключами которых являются названия товаров, а значениями — количество единиц этого товара, которые купил этот покупатель.

Напишите функцию `aggregate(*args)`, принимающую некоторое количество набор кортежей из базы данных и возвращающую сводную информацию в виде словаря, как указано выше.

Примеры см. в тесте.

<!-- grade_id: prL, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: prL_t, points: 3, grade: true, solution: false -->
```python
assert aggregate(("Ivanov","paper",10), 
("Petrov","pens",5), 
("Ivanov","marker",3), 
("Ivanov","paper",7), 
("Petrov","envelope",20), 
("Ivanov","envelope",5))=={'Ivanov': {'envelope': 5, 'marker': 3, 'paper': 17},
 'Petrov': {'envelope': 20, 'pens': 5}}
assert aggregate(("Ivanov","aaa",1), 
("Petrov","aaa",2), 
("Sidorov","aaa",3), 
("Ivanov","aaa",6), 
("Petrov","aaa",7), 
("Sidorov","aaa",8), 
("Ivanov","bbb",3), 
("Petrov","bbb",7), 
("Sidorov","aaa",345), 
("Ivanov","ccc",45), 
("Petrov","ddd",34), 
("Ziborov","eee",234), 
("Ivanov","aaa",45))=={'Ivanov': {'aaa': 52, 'bbb': 3, 'ccc': 45},
 'Petrov': {'aaa': 9, 'bbb': 7, 'ddd': 34},
 'Sidorov': {'aaa': 356},
 'Ziborov': {'eee': 234}}
assert aggregate(("TKSNUU","FKXYPUGQ",855146), 
("TKSNUU","FKXYPUGQ",930060), 
("TKSNUU","FKXYPUGQ",886973), 
("TKSNUU","FKXYPUGQ",59344), 
("TKSNUU","FKXYPUGQ",296343), 
("TKSNUU","FKXYPUGQ",193166), 
("TKSNUU","FKXYPUGQ",211696), 
("TKSNUU","FKXYPUGQ",821064), 
("TKSNUU","FKXYPUGQ",672846), 
("TKSNUU","FKXYPUGQ",820341), 
("TKSNUU","FKXYPUGQ",350693), 
("TKSNUU","FKXYPUGQ",469538), 
("TKSNUU","FKXYPUGQ",849069), 
("TKSNUU","FKXYPUGQ",502007), 
("TKSNUU","FKXYPUGQ",961595), 
("TKSNUU","FKXYPUGQ",747271), 
("TKSNUU","FKXYPUGQ",863648), 
("TKSNUU","FKXYPUGQ",952069), 
("TKSNUU","FKXYPUGQ",286019), 
("TKSNUU","FKXYPUGQ",364841), 
("TKSNUU","FKXYPUGQ",455930), 
("TKSNUU","FKXYPUGQ",100486), 
("TKSNUU","FKXYPUGQ",335026), 
("TKSNUU","FKXYPUGQ",197672), 
("TKSNUU","FKXYPUGQ",217640), 
("TKSNUU","FKXYPUGQ",612549), 
("TKSNUU","FKXYPUGQ",622501), 
("TKSNUU","FKXYPUGQ",96554), 
("TKSNUU","FKXYPUGQ",327166), 
("TKSNUU","FKXYPUGQ",425399), 
("TKSNUU","FKXYPUGQ",362309), 
("TKSNUU","FKXYPUGQ",78477), 
("TKSNUU","FKXYPUGQ",258916), 
("TKSNUU","FKXYPUGQ",297923), 
("TKSNUU","FKXYPUGQ",8891), 
("TKSNUU","FKXYPUGQ",13639), 
("TKSNUU","FKXYPUGQ",77308), 
("TKSNUU","FKXYPUGQ",707620), 
("TKSNUU","FKXYPUGQ",68205), 
("TKSNUU","FKXYPUGQ",256702), 
("TKSNUU","FKXYPUGQ",668334), 
("TKSNUU","FKXYPUGQ",968673), 
("TKSNUU","FKXYPUGQ",138125), 
("TKSNUU","FKXYPUGQ",222904), 
("TKSNUU","FKXYPUGQ",214091), 
("TKSNUU","FKXYPUGQ",500231), 
("TKSNUU","FKXYPUGQ",19611), 
("TKSNUU","FKXYPUGQ",491343), 
("TKSNUU","FKXYPUGQ",404307), 
("TKSNUU","FKXYPUGQ",68367), 
("TKSNUU","FKXYPUGQ",287107), 
("TKSNUU","FKXYPUGQ",794935), 
("TKSNUU","FKXYPUGQ",254217), 
("TKSNUU","FKXYPUGQ",206370), 
("TKSNUU","FKXYPUGQ",202761), 
("TKSNUU","FKXYPUGQ",929017), 
("TKSNUU","FKXYPUGQ",843359), 
("TKSNUU","FKXYPUGQ",955269), 
("TKSNUU","FKXYPUGQ",134139), 
("TKSNUU","FKXYPUGQ",946168), 
("TKSNUU","FKXYPUGQ",967781), 
("TKSNUU","FKXYPUGQ",856474), 
("TKSNUU","FKXYPUGQ",465070), 
("TKSNUU","FKXYPUGQ",580526), 
("TKSNUU","FKXYPUGQ",172109), 
("TKSNUU","FKXYPUGQ",191703), 
("TKSNUU","FKXYPUGQ",207916), 
("TKSNUU","FKXYPUGQ",512264), 
("TKSNUU","FKXYPUGQ",533081), 
("TKSNUU","FKXYPUGQ",577208), 
("TKSNUU","FKXYPUGQ",831389), 
("TKSNUU","FKXYPUGQ",439158), 
("TKSNUU","FKXYPUGQ",565633), 
("TKSNUU","FKXYPUGQ",452643), 
("TKSNUU","FKXYPUGQ",164426), 
("TKSNUU","FKXYPUGQ",540743), 
("TKSNUU","FKXYPUGQ",880704), 
("TKSNUU","FKXYPUGQ",868529), 
("TKSNUU","FKXYPUGQ",240742), 
("TKSNUU","FKXYPUGQ",868865), 
("TKSNUU","FKXYPUGQ",910442), 
("TKSNUU","FKXYPUGQ",146737), 
("TKSNUU","FKXYPUGQ",820984), 
("TKSNUU","FKXYPUGQ",660948), 
("TKSNUU","FKXYPUGQ",957975), 
("TKSNUU","FKXYPUGQ",135847), 
("TKSNUU","FKXYPUGQ",401865), 
("TKSNUU","FKXYPUGQ",982859), 
("TKSNUU","FKXYPUGQ",748454), 
("TKSNUU","FKXYPUGQ",354734), 
("TKSNUU","FKXYPUGQ",525638), 
("TKSNUU","FKXYPUGQ",119140), 
("TKSNUU","FKXYPUGQ",484816), 
("TKSNUU","FKXYPUGQ",616539), 
("TKSNUU","FKXYPUGQ",682553), 
("TKSNUU","FKXYPUGQ",841541), 
("TKSNUU","FKXYPUGQ",713063), 
("TKSNUU","FKXYPUGQ",433453), 
("TKSNUU","FKXYPUGQ",465340), 
("TKSNUU","FKXYPUGQ",985635))=={'TKSNUU': {'FKXYPUGQ': 49769497}}
assert aggregate()=={}
```