# Программирование на языке Python для сбора и анализа данных
*Семинары ведут: Щуров И.В., Будылин Р.Я., НИУ ВШЭ*

Данный notebook является набором задач по курсу 
«Программирование на языке Python для сбора и анализа данных» 
(НИУ ВШЭ, 2014-15). Распространяется по лицензии [CC BY-SA 
4.0](https://creativecommons.org/licenses/by-sa/4.0/deed.ru) (кроме 
задач, про которые сказано обратное) 
На [странице курса](http://math-info.hse.ru/s15/m) находятся другие материалы.

## Домашнее задание №4
За разные задачи можно получить разное число баллов. Максимум за ДЗ можно набрать 10 баллов. Вы можете решить больше задач, чем требуется, чтобы набрать 10 баллов, чтобы подстраховаться.

Для предварительной проверки задания нужно скачать данный `ipynb`-файл на свой компьютер, открыть его в IPython Notebook/Jupyter (подробнее о том, как это сделать, см. [здесь](http://math-info.hse.ru/2015-16/%D0%9F%D1%80%D0%BE%D0%B3%D1%80%D0%B0%D0%BC%D0%BC%D0%B8%D1%80%D0%BE%D0%B2%D0%B0%D0%BD%D0%B8%D0%B5_%D0%BD%D0%B0_%D1%8F%D0%B7%D1%8B%D0%BA%D0%B5_Python_%D0%B4%D0%BB%D1%8F_%D1%81%D0%B1%D0%BE%D1%80%D0%B0_%D0%B8_%D0%B0%D0%BD%D0%B0%D0%BB%D0%B8%D0%B7%D0%B0_%D0%B4%D0%B0%D0%BD%D0%BD%D1%8B%D1%85#.D0.9A.D0.B0.D0.BA_.D0.B8.D1.81.D0.BF.D0.BE.D0.BB.D1.8C.D0.B7.D0.BE.D0.B2.D0.B0.D1.82.D1.8C)) и вставить решение каждой задачи в ячейку для кода, следующую за его условием, запустить эту ячейку, а затем запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (assertion), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и задачу нужно решить заново.

Чтобы сдать ДЗ, его надо будет загрузить на сервер в виде `ipynb`-файла. Куда загрузить — мы скажем дополнительно чуть позже.

### Задача A (1 балл).
Написать функцию `times2(x)`, принимающую на вход число `x` и возвращающую это число, увеличенное в два раза.

<!-- grade_id: prA, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prA_t1, points: 1, grade: true, solution: false -->
```python
assert(times2(1) == 2)
assert(times2(2) == 4)
assert(times2(-2) == -4)
assert(times2(0) == 0)
assert(times2(100) == 200)
```

### Задача B (1 балла).
Написать функцию `is_odd(n)`, проверяющую, является ли данное число `n` нечётным. Она должна возвращать `True`, если число нечётное, и `False`, если чётное.

<!-- grade_id: prB, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prB_t1, points: 1, grade: true, solution: false -->
```python
assert(is_odd(1))
assert(not is_odd(2))
assert(not is_odd(0))
assert(is_odd(101))
```

### Задача C (1 балл).
Написать функцию `hello(n)`, выводящую на экран строчку `Hello` `n` раз подряд и ничего не возвращающую. Например, вызов `hello(3)` должен привести к выводу на экран
```
Hello
Hello
Hello
```

<!-- grade_id: prC, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prC_t1, points: 1, grade: true, solution: false -->
```python
buffer = []
passed = True
def print(*args, end = "\n", sep=" "):
    buffer.append(sep.join(map(str,args))+end)
def my_assert(q):
    global passed
    if not q:
        passed = False
def test_hello(n):
    buffer[:]=[]
    s = hello(n)
    my_assert(s == None)
    my_assert("".join(buffer).strip()==("Hello\n"*n).strip())
    buffer[:]=[]
test_hello(3)
test_hello(10)
test_hello(151)
del(print)
assert(passed)
```

### Задача D. (2 балла)
Написать функцию `num_div(n)`, вычисляющую, сколько у данного числа `n` делителей (включая 1 и само это число). Функция должна работать с целыми положительными числами.

<!-- grade_id: prD, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prD_t1, points: 2, grade: true, solution: false -->
```python
assert(num_div(1) == 1)
assert(num_div(2) == 2)
assert(num_div(3) == 2)
assert(num_div(100) == 9)
assert(num_div(1000) == 16)
assert(num_div(139) == 2)
assert(num_div(1289237) == 2)
```

### Задача E (1 балл).
Написать функцию `is_prime(n)`, проверяющую, является ли данное число `n` простым. Она должна возвращать `True`, если число является простым, и `False` в противном случае. **Примечание.** Можете использовать ранее написанную функцию `num_div(n)`. (**Вопрос:** Является ли это самым эффективным способом решения данной задачи?)

<!-- grade_id: prE, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prE_t1, points: 1, grade: true, solution: false -->
```python
assert(is_prime(2) and 
       is_prime(3) and 
       is_prime(139) and 
       is_prime(617) and 
       is_prime(293) and 
       not is_prime(12) and 
       not is_prime(13*7))
```

### Задача F. (2 балла)
Написать функцию `med3(x,y,z)`, возвращающую медиану из чисел `x`, `y`, `z` (то есть то из этих чисел, которое будет стоять посередине, если их упорядочить). Пользоваться какими-либо библиотечными функциями нельзя!

<!-- grade_id: prF, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prF_t1, points: 2, grade: true, solution: false -->
```python
from itertools import permutations
def test_med3(x,y,z,m):
    for t in permutations([x,y,z]):
        assert(med3(*t) == m)
test_med3(1,1,1,1)
test_med3(1,1,2,1)
test_med3(1,2,2,2)
test_med3(1,2,3,2)
test_med3(0,1,2,1)
test_med3(100,200,300,200)
test_med3(100,101,1000,101)
```

### Задача G. (2 балла)
Написать функцию `mean(x1,x2,…)`, принимающую на вход несколько чисел и возвращающую их среднее арифметическое.

<!-- grade_id: prG, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prG_t1, points: 2, grade: true, solution: false -->
```python
assert(mean(1) == 1)
assert(mean(1,2,3) == 2)
assert(mean(1,5, 8, 12) == 6.5)
assert(mean(8, 9, 1, 2, 3, 11, 15) == 7.)
```

### Задача H. (3 балла)
(Источник: informatics.mccme.ru.) Написать функцию `second_occurance(s)`, принимающую на вход строку `s` и возвращающую индекс **второго** вхождения символа `'f'`. Если символ `'f'` в данной строке встречается только один раз, функция должна вернуть число -1, а если не встречается ни разу, функция должна вернуть число -2.

<!-- grade_id: prH, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prH_t1, points: 3, grade: true, solution: false -->
```python
assert(second_occurance("comfort") == -1)
assert(second_occurance("coffee") == 3)
assert(second_occurance("qwerty") == -2)
assert(second_occurance("f") == -1)
assert(second_occurance("oooooooooooooooooof") == -1)
assert(second_occurance("ff") == 1)
assert(second_occurance("foooooooooooooof") == 15)
assert(second_occurance("oooooooooooooff") == 14)
assert(second_occurance("ofofofofofofofofo") == 3)
```

### Задача J. (3 балла)
Написать функцию `is_increasing(L, strict)`, проверяющую, является ли список `L` монотонно возрастающим/неубывающим. Если параметр `strict` установлен в `True`, то следует осуществить проверку на строгое возрастание (то есть функция должна возвращать `True` только если каждый следующий элемент списка строго больше предыдущего). Если параметр `strict` установлен в `False`, то функция должна проверять неубывание списка, то есть функция должна возвращать `True` если каждый следующий элемент списка не меньше предыдущего. Если параметр `strict` не установлен (функции передан только один параметр), то нужно проводить проверку на строгое возрастание.

<!-- grade_id: prJ, points: 0, grade: false, solution: true -->
```python
# Внесите решение в эту ячейку
```

<!-- grade_id: prJ_t1, points: 3, grade: true, solution: false -->
```python
assert(is_increasing([1,2,3]))
assert(is_increasing([1,2,2],False))
assert(not is_increasing([1,2,2]))
assert(is_increasing([1,2,5,10,100]))
assert(is_increasing([1,2,5,10,100],True))
assert(not is_increasing([2,1]))
assert(not is_increasing([2,1],True))
assert(not is_increasing([2,1],False))
assert(is_increasing([1,2]))
assert(is_increasing([1,2], True))
assert(is_increasing([1,2], False))
assert(is_increasing([0,0], False))
assert(not is_increasing([0,0], True))
assert(not is_increasing([1,2,3,4,5,5]))
assert(not is_increasing([1,1, 2, 3, 4, 5]))
```

```python

```